---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Python como calculadora

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

O python pode ser utilizado para operações básicas de soma, dividir e multiplicar

```
 a = 2 + 2 
 b = 5 * 2 
 print(a * b)
 40
```

`@instructions`
crie uma variável c em que o resultado é a multiplicação de ```a ```  por ```b ```

`@hint`


`@pre_exercise_code`
```{python}
a= 5
b = 3
# define c
```

`@sample_code`
```{python}
# define c

print(c)
```

`@solution`
```{python}
# define c
c = a * b
```

`@sct`
```{python}
c = 15
```
