---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Python como calculadora

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

O python pode ser utilizado para operações básicas de soma, dividir e multiplicar

```
 a = 2 + 2 
 b = 5 * 2 
 print(a * b)
 40
```

`@instructions`
crie uma variável ```c ``` em que o resultado é a multiplicação de ```a ```  por ```b ```

`@hint`
a multiplicação deverá atribuir a c o valor de a multiplicado por b

`@pre_exercise_code`
```{python}
a= 5
b = 3
```

`@sample_code`
```{python}
# define c

print(c)
```

`@solution`
```{python}
# define c
c = a * b
print(c)
```

`@sct`
```{python}
Ex().check_object("c").has_equal_value("15")
```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: dff4f65960
xp: 100
```



`@instructions`
O python pode ser

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
